# Assignment 2 - Network Connectivity

In this assignment you will go through the process of importing and analyzing an internal email communication network between employees of a mid-sized manufacturing company. 
Each node represents an employee and each directed edge between two nodes represents an individual email. The left node represents the sender and the right node represents the recipient. We will also store the timestamp of each email. 

```python
import networkx as nx

#!head assets/email_network.txt
```

### Question 1

Using networkx, load up the directed multigraph from `assets/email_network.txt`. Make sure the node names are strings.

*This function should return a directed multigraph networkx graph.*

<!-- grade_id: cell-60316968695c6f20, points: 0, grade: false, solution: true -->
```python
def answer_one():
    
    #import data as dataframe
    import pandas as pd
    email_data = 'assets/email_network.txt'
    email_df = pd.read_csv(email_data, delimiter='\t').astype(str)
    
    # create graph from dataframe;
    graph = nx.from_pandas_edgelist(email_df, source='#Sender', target='Recipient', edge_attr='time', create_using=nx.MultiDiGraph())
    
    # raise NotImplementedError()
    
    # Testing
    # print(f'Is Multigraph: {graph.is_multigraph()}')
    # print(f'Is Digraph: {graph.is_directed()}')
    # print(f'Num verts: {len(graph.nodes)}')
    # print(f'Num Edge: {len(graph.edges)}')
    # print(f'Datatype: {email_df.dtypes}')
    # print(f'Node dtypes: {type(list(graph.nodes)[0])}')
    
    return graph

# answer_one()
```

<!-- grade_id: cell-483a56c89f9db231, points: 0, grade: true, solution: false -->
```python
ans_one = answer_one()
```

### Question 2

How many employees are represented in the network?

How many `sender`->`recepient` pairs of employees are there in the network such that `sender` sent at least one email to `recepient`? Note that even if a `sender` sent multiple messages to a `recepient`, they should only be counted once. You should exclude cases where an employee sent emails to themselves from this count. 

*This function should return a tuple with two integers (#employees, # `sender`->`recepient` pairs).*

<!-- grade_id: cell-16a78f22207d9f95, points: 0, grade: false, solution: true -->
```python
def answer_two():
    
    graph = answer_one()
    
    # number of employees
    # number of vertices in the graph
    num_employees = len(list(graph.nodes))
    
    
#     # number of sender/recipient pairs
#     # method 1: count manually
    
#     unique_connections = [e for e in graph.edges if e[0] != e[1]]
#     num_pairs = len(unique_connections)
    
    
    # number of sender/recipient pairs
    # method 2: convert to other type of graph
    
    # graph = nx.Graph(graph)
    # num_pairs_2 = len(graph.edges)
    
    graph = nx.DiGraph(graph)
    num_pairs_2 = len(graph.edges)
    
    # Testing
    # print(f'Is Multigraph: {graph.is_multigraph()}')
    # print(f'Is Digraph: {graph.is_directed()}')
    # print(f'Num verts: {len(graph.nodes)}')
    # print(f'Num Edge: {len(graph.edges)}')
    # print(f'Datatype: {email_df.dtypes}')
    # print(f'Node dtypes: {type(list(graph.nodes)[0])}')
    
#     # verification: same answer using both methods?
#     print(f'Method 1 num connections: {num_pairs}')
#     print(f'Method 2 num connections: {num_pairs_2}')
#     print(f'Consistent? {num_pairs == num_pairs_2}')
    


    # raise NotImplementedError()
    return (num_employees, num_pairs_2)

answer_two()
```

Output:
```
Is Multigraph: False
Is Digraph: True
Num verts: 167
Num Edge: 5784
(167, 5784)
```

<!-- grade_id: cell-5b6391549b076d2b, points: 1, grade: true, solution: false -->
```python
ans_two = answer_two()

```

### Question 3

* Part 1. Assume that information in this company can only be exchanged through email.

    When an employee sends an email to another employee, a communication channel has been created, allowing the sender to provide information to the reciever, but not viceversa. 

    Based on the emails sent in the data, is it possible for information to go from every employee to every other employee?


* Part 2. Now assume that a communication channel established by an email allows information to be exchanged both ways. 

    Based on the emails sent in the data, is it possible for information to go from every employee to every other employee?


*This function should return a tuple of bools (part1, part2).*

<!-- grade_id: cell-606dd37fb2f55358, points: 0, grade: false, solution: true -->
```python
def answer_three():
    
    # MultiDiGraph
    # test if there are multiple edges
    graph = answer_one()

    
    # test to see if strongly connected
    # print('Is strongly connected:', nx.is_strongly_connected(graph))
    # print('Is weakly connected:', nx.is_weakly_connected(graph))
    
    # raise NotImplementedError()
    return (nx.is_strongly_connected(graph), nx.is_weakly_connected(graph))

# answer_three()
```

Output:
```
(False, True)
```

<!-- grade_id: cell-82b3f0bc45e2895f, points: 1, grade: true, solution: false -->
```python
ans_three = answer_three()

```

### Question 4

How many nodes are in the largest weakly connected component of the graph?

*This function should return an int.*

<!-- grade_id: cell-39196f4c27b35534, points: 0, grade: false, solution: true -->
```python
def answer_four():
    
    graph = answer_one()
    
    # raise NotImplementedError()
    list(nx.strongly_connected_components(graph))
    
    return len(max(nx.weakly_connected_components(graph)))

# answer_four()
```

Output:
```
167
```

<!-- grade_id: cell-2b1b7b06ecfa751d, points: 1, grade: true, solution: false -->
```python
ans_four = answer_four()
```

### Question 5

How many nodes are in the largest strongly connected component?

*This function should return an int*

<!-- grade_id: cell-879d17a9db5c3790, points: 0, grade: false, solution: true -->
```python
def answer_five():
    
    graph = answer_one()
    
    largest = max(nx.strongly_connected_components(graph), key=len)
    # raise NotImplementedError()
    
    return len(largest)
# answer_five()
```

<!-- grade_id: cell-b0524f7dc1fbdec4, points: 1, grade: true, solution: false -->
```python
ans_five = answer_five()

```

### Question 6

Using the NetworkX functions `strongly_connected_components` and `subgraph`, find the subgraph of nodes in the largest strongly connected component. 
Call this graph G_sc.

*This function should return a networkx MultiDiGraph named G_sc.*

<!-- grade_id: cell-a81e6f88ef84578c, points: 0, grade: false, solution: true -->
```python
def answer_six():
    graph = answer_one()
    
    max_scc = max(nx.strongly_connected_components(graph), key=len)
    
    G_sc = nx.subgraph(graph, max_scc)
    # raise NotImplementedError()
    
    # test
    print(f'Is strongly connected: {nx.is_strongly_connected(G_sc)}')
    print(f'Is Multigraph: {G_sc.is_multigraph()}')
    print(f'Is directed: {G_sc.is_directed()}')
    
    return G_sc
# answer_six()
```

Output:
```
Is Multigraph: True
Is Digraph: True
Num verts: 167
Num Edge: 82927
Is strongly connected: True
Is Multigraph: True
Is directed: True
<networkx.classes.multidigraph.MultiDiGraph at 0x7fcd91a0b4c0>
```

<!-- grade_id: cell-cf148ef273b3b19c, points: 0, grade: true, solution: false -->
```python
ans_six = answer_six()
assert type(ans_six) == nx.MultiDiGraph , "Your return type should be a MultiDiGraph object"

```

### Question 7

What is the average distance between nodes in G_sc?

*This function should return a float.*

<!-- grade_id: cell-1a3075039b3c2c0b, points: 0, grade: false, solution: true -->
```python
def answer_seven():
    
    graph = answer_six()
    
    # raise NotImplementedError()
    
    return nx.average_shortest_path_length(graph)

# answer_seven()
```

<!-- grade_id: cell-5b374fdd48f37e02, points: 1, grade: true, solution: false -->
```python
ans_seven = answer_seven()

```

### Question 8

What is the largest possible distance between two employees in G_sc?

*This function should return an int.*

<!-- grade_id: cell-4fbb8d34d00f675c, points: 0, grade: false, solution: true -->
```python
def answer_eight():
    graph = answer_six()
    
    # raise NotImplementedError()
    
    return nx.diameter(graph)

# answer_eight()
```

<!-- grade_id: cell-c5714787854ef644, points: 1, grade: true, solution: false -->
```python
ans_eight = answer_eight()

```

### Question 9

What is the set of nodes in G_sc with eccentricity equal to the diameter?

*This function should return a set of the node(s).*

<!-- grade_id: cell-bcfab850744277ed, points: 0, grade: false, solution: true -->
```python
def answer_nine():
    graph = answer_six()
    
    # raise NotImplementedError()
    
    nodes = nx.periphery(graph)
    
    # test, comment out
    # nodes = nx.center(graph)
    
    return set(nodes)

# answer_nine()
```

<!-- grade_id: cell-77c9ca0b94df3d6f, points: 1, grade: true, solution: false -->
```python
ans_nine = answer_nine()
assert type(ans_nine) == set, "Student answer must return a set"
```

### Question 10

What is the set of node(s) in G_sc with eccentricity equal to the radius?

*This function should return a set of the node(s).*

<!-- grade_id: cell-fd3d43ea37da03f7, points: 0, grade: false, solution: true -->
```python
def answer_ten():
    
    graph = answer_six()

    
    # nodes = nx.periphery(graph)
        
    nodes = nx.center(graph)    
    # raise NotImplementedError()

    
    return set(nodes)

# answer_ten()
    
```

<!-- grade_id: cell-bfd2ee304bc25264, points: 1, grade: true, solution: false -->
```python
ans_ten = answer_ten()
assert type(ans_ten) == set, "Student answer must return a set"
```

### Question 11

Which node in G_sc has the most shortest paths to other nodes whose distance equal the diameter of G_sc?


For the node with the most such shortest paths, how many of these paths are there?


*This function should return a tuple (name of node, number of paths).*

<!-- grade_id: cell-c40858baf3d92139, points: 0, grade: false, solution: true -->
```python
def answer_eleven():

    graph = answer_six()
    
    periphery = nx.periphery(graph)
    diam = nx.diameter(graph)
    
    
    max_vert = max(periphery, key=lambda x: nx.descendants_at_distance(graph, x, diam))
    path_lengths = [len(nx.descendants_at_distance(graph, vert, diam)) for vert in periphery]
    
    
    # Test
    # print(f'Is strongly connected: {nx.is_strongly_connected(graph)}')
    # print(f'Is Multigraph: {graph.is_multigraph()}')
    # print(f'Is directed: {graph.is_directed()}')
    # print(f'Periphery set: {periphery}')
    # print(f'Path length list: {path_lengths}')
    
    
    # raise NotImplementedError()
    return (max_vert, max(path_lengths))

# answer_eleven()
```

Output:
```
(97, 63)
```

<!-- grade_id: cell-f79b06650f61cf37, points: 1, grade: true, solution: false -->
```python
ans_eleven = answer_eleven()
assert type(ans_eleven) == tuple, "Student answer must be a tuple"
```

### Question 12

Suppose you want to prevent communication flow from the node that you found in question 11 to node 10. What is the smallest number of nodes you would need to remove from the graph (you're not allowed to remove the node from the previous question or 10)? 

*This function should return an integer.*

<!-- grade_id: cell-a4606db8971dcbba, points: 0, grade: false, solution: true -->
```python
def answer_twelve():

    graph = answer_six()
    source_vert = answer_eleven()[0]
    
    node_cut = nx.minimum_node_cut(graph, source_vert, 10)
    
    # raise NotImplementedError()
    
    
    return len(node_cut)
# answer_twelve()
```

Output:
```
1
```

<!-- grade_id: cell-509cfa9f4136124d, points: 1, grade: true, solution: false -->
```python
ans_twelve = answer_twelve()

```

### Question 13

Convert the graph G_sc into an undirected graph by removing the direction of the edges of G_sc. Call the new graph G_un. 


*This function should return a networkx Graph.*

<!-- grade_id: cell-612b40425f0851e8, points: 0, grade: false, solution: true -->
```python
def answer_thirteen():
    
    graph = answer_six()
    
    
    # idea 1: return Multigraph (undirected)
    G_un = graph.to_undirected()
    
    
    # idea 2: return Graph (not multigraph)
    # comment this out to return NultiGraph
    G_un = nx.Graph(G_un)
    
    
    
    # Test
    # print(f'Is strongly connected: {nx.is_strongly_connected(G_un)}')
    # print(f'Is Multigraph: {G_un.is_multigraph()}')
    # print(f'Is directed: {G_un.is_directed()}')
    # print(f'Periphery set: {periphery}')
    # print(f'Path length list: {path_lengths}')
    # raise NotImplementedError()
    
    return G_un

# type(answer_thirteen())
```

Output:
```
networkx.classes.graph.Graph
```

<!-- grade_id: cell-d1c0627a327cd774, points: 0, grade: true, solution: false -->
```python
ans_thirteen = answer_thirteen()
assert type(ans_thirteen) == nx.Graph , "Your return type should be a Graph object"

```

### Question 14

What is the transitivity and average clustering coefficient of graph G_un?

*This function should return a tuple (transitivity, avg clutering).*     
*Note: DO NOT round up your answer.*

<!-- grade_id: cell-63e48b1eed1a1fb8, points: 0, grade: false, solution: true -->
```python
def answer_fourteen():
    
    G_un = answer_thirteen()
    
    trans = nx.transitivity(G_un)
    
    av_clust = nx.average_clustering(G_un)
    
    # raise NotImplementedError()
    
    return (trans, av_clust)

# answer_fourteen()
```

Output:
```
(0.570111160700385, 0.6975272437231418)
```

<!-- grade_id: cell-41dda91202f58e7e, points: 1, grade: true, solution: false -->
```python
ans_fourteen = answer_fourteen()
assert type(ans_fourteen) == tuple, "Student answer must be a tuple"

```